<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2026/27
Prof. Dr. Heiner Giefers

**Übung 1: Welche Funktion hat der Sekundärspeicher in einem Computer?**

a\) Alle Berechnungen und die Logik des Programms ausführen\
b\) Abruf von Webseiten über das Internet\
c\) Informationen langfristig zu speichern, auch über einen Stromausfall hinaus\
d\) Eingaben des Benutzers entgegennehmen

<details>
<summary><b>Antwort anzeigen</b></summary>

**Antwort: c) Informationen langfristig zu speichern, auch über einen Stromausfall hinaus**

Der Sekundärspeicher speichert Informationen langfristig. Die gespeicherten Informationen bleiben auch nach einem Stromausfall erhalten.

</details>

**Aufgabe 4: Was enthält Maschinencode?**

a\) Der Python-Interpreter\
b\) Die Tastatur\
c\) Die Python-Quelldatei\
d\) Ein Textverarbeitungsdokument

<details>
<summary><b>Antwort anzeigen</b></summary>

**Antwort: a) Der Python-Interpreter**

Der Python-Interpreter enthält Maschinencode.

</details>

**Aufgabe 5: Was ist falsch an folgendem Code? Verbessern Sie ihn in der folgenden Zelle.**

```python
>>> primt 'Hello world!'
File “<stdin>“, line 1
  primt 'Hello world!'
                     ^
SyntaxError: invalid syntax
```

<!-- grade_id: cell-41c9bb1915328d6e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
%%capture cap --no-stderr --no-display
c = _i
exec(c)
```

<!-- grade_id: cell-395e4adf2fbca228, points: 1, grade: true, solution: false -->
```python
assert cap.stdout.strip() == "Hello world!", f"Die Ausgabe {cap.stdout} stimmt nicht."
```

**Übung 6: Wo im Computer wird eine Variable wie `x` gespeichert, nachdem die folgende Python-Zeile beendet ist?**

```python
x = 123
```

a\) CPU\
b) Hauptspeicher\
c) Sekundärspeicher\
d) Eingabegeräte\
e) Ausgabegeräte

<details>
<summary><b>Antwort anzeigen</b></summary>

**Antwort: b) Hauptspeicher**

Während das Programm läuft, werden Variablen im Hauptspeicher gespeichert.

</details>

**Übung 7: Was wird das folgende Programm ausgeben:**

```python
x = 43
x = x + 1
print(x)
```

a\) 43\
b) 44\
c) x + 1\
d) Fehler, denn x = x + 1 ist mathematisch nicht möglich

<details>
<summary><b>Antwort anzeigen</b></summary>

**Antwort: b) 44**

Die zweite Zuweisung erhöht den bisherigen Wert von `x` um 1. Deshalb gibt `print(x)` den Wert 44 aus.

</details>