<figure>
  <IMG SRC="https://raw.githubusercontent.com/fhswf/datasets/refs/heads/main/images/fh-swf-logo.svg" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2026/27
Prof. Dr. Heiner Giefers

**Aufgabe 1: Schreiben Sie eine while-Schleife, die beim letzten Zeichen in der Zeichenkette beginnt und sich rückwärts bis zum ersten Zeichen in der Zeichenkette vorarbeitet, wobei jedes Zeichen in einer eigenen Zeile ausgegeben wird, natürlich rückwärts.**

<!-- grade_id: cell-9f6dc7ff3dbe0557, points: 0, grade: false, solution: true -->
```python
def drucke_zeichen_rueckwaerts(zeichenkette):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
%%capture cap --no-stderr --no-display
drucke_zeichen_rueckwaerts("Hello")
```

<!-- grade_id: cell-a2fe2825342aaa23, points: 1, grade: true, solution: false -->
```python
"""Test der Funktion drucke_zeichen_rueckwaerts."""
from nose.tools import assert_equal
assert_equal(str(cap), 'o\nl\nl\ne\nH\n', "Die Ausgabe stimmt nicht")
```

**Aufgabe 2: Schreiben Sie eine Funktion `extrahiere_zeichen`, die aus einer Zeichenkette den Teil nach dem Doppelpunkt extrahiert:**

**`zeichenkette = 'X-DSPAM-Confidence:0.8475'`**

**Verwenden Sie `find` und String-Slicing, um den Teil nach dem Doppelpunkt zu extrahieren, und verwenden Sie anschließend `float`, um ihn in eine Fließkommazahl zu konvertieren.**

<!-- grade_id: cell-997127479ac5692f, points: 0, grade: false, solution: true -->
```python
def extrahiere_zeichen(zeichenkette):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a2390f7eb31779fa, points: 1, grade: true, solution: false -->
```python
"""Test der Funktion extrahiere_zeichen."""
from nose.tools import assert_equal
assert_equal(0.8475, extrahiere_zeichen("X-DSPAM-Confidence:0.8475"))
assert_equal(0.5, extrahiere_zeichen("X-DSPAM-Confidence:0.5"))
```

**Aufgabe 3: Schreiben Sie eine Funktion `count_vocals`, welche die fünf Vokale `a`, `e`, `i`, `o` und `u` in einem als Parameter übergebenen Text zählt und die Summe aller Vokale zurückgibt. Groß- und Kleinschreibung sollen dabei gleich behandelt werden.**

<!-- grade_id: cell-6405df4aa3f9f2b7, points: 0, grade: false, solution: true -->
```python
def count_vocals(text):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8012e60ca6ad29a0, points: 1, grade: true, solution: false -->
```python
"""Test der Funktion count_vocals."""
from nose.tools import assert_equal
assert_equal(0, count_vocals(""))
# Beispiel - Text Quelle: https://de.wikipedia.org/wiki/Vokal
assert_equal(58, count_vocals("ein vokal oder selbstlaut ist ein sprachlaut, bei dessen artikulation der phonationsstrom weitgehend ungehindert durch den mund ausströmen kann. vokale sind im allgemeinen stimmhaft."))
assert_equal(58, count_vocals("Ein Vokal oder Selbstlaut ist ein Sprachlaut, bei dessen Artikulation der Phonationsstrom weitgehend ungehindert durch den Mund ausströmen kann. Vokale sind im Allgemeinen stimmhaft."))
```