<figure>
  <IMG SRC="https://raw.githubusercontent.com/fhswf/datasets/refs/heads/main/images/fh-swf-logo.svg" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2026/27
Prof. Dr. Heiner Giefers

**Übung 1: Schreiben Sie eine Funktion `lese_datei`, die eine Datei einliest und den Inhalt der Datei in Großbuchstaben (Zeile für Zeile) ausgibt.**

<!-- grade_id: cell-c1e55762bc63dc93, points: 0, grade: false, solution: true -->
```python
def lese_datei(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
%%capture cap --no-stderr --no-display
with open('myfile.txt', "w") as myfile:
    myfile.write("Ein neuer\nText.\nEnde")
lese_datei("myfile.txt")
```

<!-- grade_id: cell-0aa5a4884c58af6f, points: 1, grade: true, solution: false -->
```python
assert str(cap) == 'EIN NEUER\nTEXT.\nENDE', "Die Ausgabe stimmt nicht"
```

**Übung 2: Schreiben Sie eine Funktion `berechne_durchschnitt`, die eine Datei einliest und in der Datei nach Zeilen der Form 
`X-DSPAM-Confidence: 0.8475` sucht. Die Zahl kann sich dabei ändern, die anderen Zeichen nicht. Insbesondere muss die Zeile immer mit `X-DSPAM...` beginnen.<br>Extrahieren Sie die Fließkommerzahl indem Sie die Zeile auftrennen. Berechnen Sie die Summe der Fließkommerzahlen und geben Sie den durchschnittlichen Wert der Fließkommerzahl zurück.**

<!-- grade_id: cell-8e47dfd8aabc8a20, points: 0, grade: false, solution: true -->
```python
def berechne_durchschnitt(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4ead248e87efeb5f, points: 2, grade: true, solution: false -->
```python
with open("test_dspam.txt", "w") as file:
    file.write(
        "Header\n"
        "X-DSPAM-Confidence: 0.5\n"
        "X-DSPAM-Confidence:0.9\n"
        "Footer\n"
    )

assert abs(berechne_durchschnitt("test_dspam.txt") - 0.7) < 1e-7
```