<figure>
  <IMG SRC="https://raw.githubusercontent.com/fhswf/datasets/refs/heads/main/images/fh-swf-logo.svg" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2026/27
Prof. Dr. Heiner Giefers

# Python Basics
Rückblick auf die Inhalte der letzten Kapitel:
- Schleifen und Funktionen
- Zeichenketten
- Dateien lesen und schreiben

### Schleifen und Funktionen

So gut wie jedes Programm verwendet Schleifen und Funktionen.
Es sind, neben der *bedingten Ausführung* (bzw. der *alternativen Ausführung*) die wichtigsten Elemente zur Strukturierung von Programmen.
In den kommenden Aufgaben, wollen wir noch einmal das Schreiben von Funktionen und Schleifen üben.

Funktionen in Python werden über das Schlüsselwort def definiert. Die Syntax einer Funktions-Definition sieht folgendermaßen aus:

```python
def myfunc(arg1, arg2,... argN):  
  '''Dokumentation'''  

  #Programmcode, optional mit return Anweisungen
```

Hier wird die Funktion `myfunc` definiert, welche mit den Parametern `arg1,arg2,....argN` aufgerufen wird.
Wir können bereits bei der Definition der Funktion Werte für einzelne Parameter angeben.
Jemand, der die Funktion verwendet, kann dann den Parameter weglassen und es wird der *default* Wert übernommen.
Die Parameter, die zwingend beim Aufruf der Funktion angegeben werden müssen, stehen vorne in der Liste, die mit *default*-Werten stehen hinten.

Wenn beim Programmaufruf Parameter angegeben werden, müssen diese in der gleichen Reihenfolge wie bei der Funktionsdefinition angegeben werden.
Sollen einige der *default*-Parameter übernommen werden, *fehlen* dieses Werte beim Funktionsaufruf und die Reihenfolge stimmt ab dem ersten übernommenen *default*-Parameter nicht mehr. In diesem Fall ist es notwendig, die folgenden Parameter, deren *default*-Werte überschrieben werden sollen, namentlich zu benennen.

Schauen wir uns die folgende Definition der Funktion `operator_tabelle` an.
Die Funktion besitzt drei Parameter:
`op` soll vom Typ `str` sein und die zu verwendende Operation angeben, `von` und `bis` legen den Start- bzw. Endwert der Operationstabelle fest.
Die Angabe eines Wertes für `op` ist von der Funktion gefordert. 
Für `von` und `bis` sind jeweils Werte vorgegeben, die aber nur verwendet werden, wenn entsprechende Argumente beim Funktionsaufruf weggelassen werden.

Was die Funktion tut und wie sie aufgerufen werden soll, wird auch im **`Docstring`** erklärt.
Dieser String kann am Anfang einer Funktion gesetzt werden.
Er dient zum Einen der besseren Verständlichkeit des Quellcodes.
Zum Anderen wird dieser String angezeigt, wenn die Python Hilfe-Funktion `help` mit dem Namen der Funktion aufgerufen wird.
Dies hilft den Personen, die die Funktion verwenden wollen, aber nicht unmittelbaren Zugriff auf den Quellcode haben.

**`return`** Anweisungen können grundsätzlich überall in der Funktion verwendet werden.
Es kann auch an mehreren Stellen ein `return`-Aufruf stehen, allerdings kann eine Funktion nur jeweils eine dieser Anweisungen ausführen.
Nach einem `return` wird die Funktion verlassen und bei einem erneuten Aufruf beginnt die Funktion wieder *von Vorne*.

```python
def operator_tabelle(op, von=1, bis=10):
    '''Die Funktion berechnet Additions- und Multiplikationstabellen.
    
    Parameters
    ----------
    op : str
            Der zu verwendende Operator, angegeben als String.
            Nur '+' und '*' werden unterstützt.

    von, bis: int, optional
            Start- und Endwert der Operationstabelle.
            
    Raises
    ------
    ValueError
            Wenn ein nicht unterstützter Operator angegeben wird.
    '''
    
    tabelle = ""
    if op=='+' or op=='*':
        for zeile in range(von, bis+1):
            for spalte in range(von, bis+1):
                if op=='+':
                    tabelle += f"{zeile+spalte:4d}"
                else:
                    tabelle += f"{zeile*spalte:4d}"
            tabelle += '\n'
        return tabelle
    else:
        raise ValueError("Fehler: Ich kann nur '+' oder '*' rechnen!")

s = operator_tabelle('*', bis=5)
print(s)
```

```python
help(operator_tabelle)
```

**Aufgabe 1:** Überlegen Sie zuerst, was die folgenden Funktionsaufrufe zurückgeben, und tragen Sie Ihre Vorhersage in die Variablen `a`, `b`, `c` und `d` ein. Prüfen Sie danach, ob Ihre Vorhersage stimmt.

**Hinweis:** Was passiert, wenn eine Funktion kein `return` ausführt?

```python
def f(x, y=2):
    return x * y

def g(x):
    if x > 0:
        return "positiv"
    elif x < 0:
        return "negativ"

def h(x):
    x + 1
```

<!-- grade_id: cell-0c0af323b76743c1, points: 0, grade: false, solution: true -->
```python
# Tragen Sie Ihre Vorhersage ein, ohne die Funktionen aufzurufen
a = None  # Ergebnis von f(3)
b = None  # Ergebnis von f(3, y=5)
c = None  # Ergebnis von g(0)
d = None  # Ergebnis von h(4)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ac6953b778194b46, points: 1, grade: true, solution: false -->
```python
### TEST ZELLE
assert a == f(3), 'Vorhersage für f(3) stimmt nicht'
assert b == f(3, y=5), 'Vorhersage für f(3, y=5) stimmt nicht'
assert c == g(0), 'Vorhersage für g(0) stimmt nicht'
assert d == h(4), 'Vorhersage für h(4) stimmt nicht'
```

**Aufgabe 2:** Schreiben Sie eine Funktion `fak`, die eine ganze Zahl als Parameter entgegennimmt und die Fakultät dieser Zahl als Ergebnis zurückliefert. Es gilt $n! = 1 \cdot 2 \cdot \ldots \cdot n$ und $0! = 1$. Für die Parameter soll weiterhin gelten:
- Wenn der Parameter weggelassen wird, soll standardmäßig die Fakultät von 42 berechnet werden.
- Für negative Zahlen soll die Funktion einen `ValueError` auslösen, da die Fakultät für sie nicht definiert ist.

<!-- grade_id: cell-b5428bb39a5f43e0, points: 0, grade: false, solution: true -->
```python
def fak(zahl=42):
    '''Berechnet die Fakultät einer Zahl'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b1051632392f4839, points: 1, grade: true, solution: false -->
```python
### TEST ZELLE
assert fak(5)==120, "Bei der Fakultäts-Funktion stimmt etwas nicht"
assert fak(0)==1, "Es gilt 0! = 1"
assert fak()==1405006117752879898543142606244511569936384000000000, "Default Parameter verwenden"

try:
    etype=False
    fak(-1)
except ValueError:
    etype = True
except Exception as e:
    pass
assert etype, "Bei negativem Parameter einen ValueError erzeugen"
```

**Aufgabe 3:** Schreiben Sie eine Funktion `ggt(a,b)`, die den größten gemeinsamen Teiler (ggT) zweier **positiver ganzer Zahlen** berechnet.


Hinweise: Eine effiziente Möglichkeit, den ggT zu berechnen ist der euklidische Algorithmus. Das Verfahren funktioniert folgendermaßen:

1. Solange `b!=0` (*b ungleich 0*) gilt, wiederhole folgende Schritte:
    - Berechne `r` als den ganzzahligen **Rest** der Division  `a/b`  (also: "r gleich a modulo b")
    - Setze `a=b` und `b=r`
1. Gebe a als Ergebnis zurück

Rechnen Sie zunächst den ggT von `48` und `18` von Hand nach diesem Verfahren durch. Notieren Sie dazu für jeden Durchlauf die Werte von `a`, `b` und `r`. Schreiben Sie erst danach das Programm. Das Ergebnis muss `6` sein.

<!-- grade_id: cell-5d949f7a0cf44eb5, points: 0, grade: false, solution: true -->
```python
def ggt(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d0067e7013d54fce, points: 1, grade: true, solution: false -->
```python
### TEST ZELLE
test_cases = [(48, 18, 6),(123, 321, 3),(545,654, 109),(459,666, 9),(7, 13, 1)]
assert all([ggt(a,b)==c for a, b, c in test_cases]), ''
```

**Zusatzaufgabe:**

Sie haben nun eine `ggt` Funktion, die den größten gemeinsamen Teiler zweier Zahlen berechnet.
Falls Sie den ggT von mehr als 2 Zahlen berechnen wollen, können Sie die Funktion `ggt` mehrfach benutzen.
Nehmen wir an, Sie wollen den ggT der Zahlen `a`, `b` und `c` berechnen.
Der ggT ist kommutativ und assoziativ. Für drei Zahlen gilt zum Beispiel:

$ggT(a,b,c) = ggT(ggT(a,b),c) = ggT(a,ggT(b,c))$.

Die unterschiedliche Klammerung folgt aus der Assoziativität; dass Zahlen vertauscht werden können, aus der Kommutativität, also $ggT(a,b) = ggT(b,a)$. Daher kann die Funktion `ggt` mehrfach aufgerufen werden, um den ggT mehrerer Zahlen zu berechnen.

**Schreiben Sie eine Funktion `ggt_x`, die den ggT von 2 und mehr Zahlen berechnet.**

**Hinweise:** Verwenden Sie variable (Positions-) Parameter nach dem Muster `ggt_x(a, b, *weitere)`.

<!-- grade_id: cell-761c8e4e22574cad, points: 0, grade: false, solution: true -->
```python
def ggt_x(a, b, *weitere):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-655115bb419f4d2e, points: 1, grade: true, solution: false -->
```python
### TEST ZELLE
assert ggt_x(120, 540) == 60, 'Die Funktion soll auch mit 2 Parametern funktionieren.'
assert ggt_x(120, 540, 1200, 3420, 5620) == 20, 'Bei der Lösung stimmt etwas nicht.'
```

### Strings und Konvertierung

**Aufgabe 4:** Schreiben Sie eine Funktion `umdrehen`, die einen String als Parameter entgegennimmt und den String in umgekehrter Reihenfolge zurückgibt. Also z. B.:
```python
umdrehen("Hallo")   # soll "ollaH" zurückgeben
umdrehen("abc")     # soll "cba" zurückgeben
umdrehen("")        # soll "" zurückgeben
```

**Hinweise:** Starten Sie mit einem leeren String `umgekehrt = ""` und gehen Sie mit einer `for`-Schleife Zeichen für Zeichen durch den Text. Setzen Sie jedes Zeichen **vor** den bisherigen Inhalt von `umgekehrt`. Testen Sie die Schleife zuerst mit einem kurzen Wort, z. B. "abc", und geben Sie `umgekehrt` mit `print` aus.

*Ausblick:* Später lernen Sie mit *Slicing* eine deutlich kürzere Möglichkeit kennen, einen String umzudrehen.

<!-- grade_id: cell-9a7450e918c04bfc, points: 0, grade: false, solution: true -->
```python
def umdrehen(text):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-bd06eef113e94428, points: 1, grade: true, solution: false -->
```python
### TEST ZELLE
assert umdrehen('Hallo') == 'ollaH', 'Aus "Hallo" soll "ollaH" werden'
assert umdrehen('abc') == 'cba', 'Aus "abc" soll "cba" werden'
assert umdrehen('a') == 'a', 'Ein einzelnes Zeichen bleibt gleich'
assert umdrehen('') == '', 'Der leere String bleibt leer'
```

### Dateien lesen und Schreiben

Das Lesen und Schreiben von Daten aus bzw. in Dateien ist in vielen Anwendungen notwendig. In Python kann man Dateien mit der Funktion `open` öffnen. Die Funktion liefert ein *File-Objekt* zurück, über das man auf die Datei zugreifen kann. Mit `close` wird die Datei wieder geschlossen.
Zeichenketten können mit der Methode `write` in die Datei geschrieben werden, mit `read` oder `readlines` liest man den Inhalt der Datei wieder ein.

**Aufgabe 5a:** Schreiben Sie drei oder mehr Strings als Zeilen in die Datei `ausgabe.txt` und schließen Sie die Datei anschließend.

**Hinweis:** Verwenden Sie die `with`-Anweisung, um die Datei zu öffnen. Dann brauchen Sie sich nicht um das Schließen der Datei zu kümmern.

<!-- grade_id: cell-de17d9f0cbc44cd5, points: 0, grade: false, solution: true -->
```python
file_name = "ausgabe.txt"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-69e6ab81ccea42e3, points: 1, grade: true, solution: false -->
```python
### TEST ZELLE
try:
    with open(file_name, "r") as file:
        assert file.read().count('\n')>2, '"ausgabe.txt" enthält weniger als 3 Zeilen'
except FileNotFoundError:
    print('"ausgabe.txt" ist nicht erstellt')
    raise
```

**Aufgabe 5b:** Öffnen Sie die Datei `ausgabe.txt` zum Lesen und geben Sie den Inhalt aus. Schließen Sie die Datei anschließend.

<!-- grade_id: cell-75df4ee58f9e4a43, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2a7ff9e21c0543f4, points: 1, grade: true, solution: false -->
```python
### TEST ZELLE
from unittest.mock import patch
with patch('__main__.print') as mock_print:
     exec(_i)
mock_print.assert_called()
```

**Hinweis:** Da bei der Arbeit mit Dateien Fehler auftreten können, ist die Fehlerbehandlung wichtig. Mit der `with`-Anweisung können Sie sicherstellen, dass Dateien auch bei auftretenden Fehlern ordnungsgemäß geschlossen werden. `with` erzeugt ein sogenanntes *Context-Manager-Objekt*, das die Datei beim Verlassen des `with`-Blocks automatisch schließt. Beispiel:
```python
with open("ausgabe.txt", "r") as f:
    nummer = 1
    for line in f.readlines():
        print(f"{nummer}: {line}", end="")
        nummer += 1
```