<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2021/22
Prof. Dr. Heiner Giefers / Prof. Dr. Doga Arinir

**Übung 1:** Das Programm aus Kapitel 9, Übung 3 soll auf folgende Weise überarbeitet werden:

Die <code>From&nbsp;</code>-Zeilen sollen gelesen und geparst werden, um im Folgenden die Adressen aus der Zeile zu extrahieren. Dabei soll mit Hilfe eines Wörterbuchs die Anzahl der Nachrichten von jeder Person gezählt werden.

Nachdem alle Daten gelesen wurden, soll die Person mit den meisten Nachrichten ausgegeben werden, indem eine Liste von (count, email)-Tupeln aus dem Wörterbuch erstellt wird. Dann muss die Liste in umgekehrter Reihenfolge sortiert und schließlich die Person mit den meisten Nachrichten ausgegeben werden.

<!-- grade_id: cell-88a0e8684be45c0a, points: 0, grade: false, solution: true -->
```python
def max_count_mail_adresses(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-871ba8cd03698f18, points: 1, grade: true, solution: false -->
```python
"""max_count_mails Tests"""
from nose.tools import assert_equal

fname = "files/mbox-short.mbox"
res = max_count_mail_adresses(fname)

assert_equal("cwen@iupui.edu", res[0], "Die E-Mail Adresse ist falsch")
assert_equal(5, res[1], "Die E-Mail Anzahl ist falsch")
```

**Übung 2:** Schreiben Sie ein Programm, welches die Häufigkeit von Nachrichten pro Tageszeit ermittelt. Die Tageszeit soll in Intervalle von jeweils einer Stunde eingeteilt werden. Es genügt also, den Stunden-Anteil des Zeitstempels zu betrachten.

Die Stunden sollen aus der <code>From&nbsp;</code>-Zeile extrahiert werden, indem die Uhrzeit-Zeichenfolge gefunden und diese Zeichenfolge dann anhand des Doppelpunkts in Teile zergliedert wird. Sobald die Nachrichtenhäufigkeit für jede Stunde gesammelt wurde, sollen diese als Dict zurückgegeben werden.

<!-- grade_id: cell-c0448436a6364f59, points: 0, grade: false, solution: true -->
```python
def count_mail_per_hour(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4ad69e362fb2d53a, points: 1, grade: true, solution: false -->
```python
"""max_count_mails Tests"""
from nose.tools import assert_equal

fname = "files/mbox-short.mbox"
res = count_mail_per_hour(fname)

assert_equal(12, len(res), "Die Uhrzeit-Anzahl ist falsch")
assert_equal([('04', 3), ('06', 1), ('07', 1), ('09', 2), ('10', 3), ('11', 6), ('14', 1), ('15', 2), ('16', 4), ('17', 2), ('18', 1), ('19', 1)],
             res, "Die Anzahlen für die Uhrzeiten sind falsch")

```

**Übung 3:** Schreiben Sie ein Programm, das eine Datei liest und die Buchstaben in absteigender Reihenfolge der Häufigkeit ausgibt. Das Programm sollte alle Eingaben in Kleinbuchstaben umwandeln und nur die Buchstaben `a`-`z` zählen. Es soll keine Leerzeichen, Ziffern, Interpunktionszeichen oder irgendetwas anderes als die Buchstaben  `a`-`z` zählen.

Dann sollen Textbeispiele aus verschiedenen Sprachen als Eingabe dienen und analyisiert werden, wie die Buchstabenhäufigkeit zwischen den Sprachen variiert. Vergleichen Sie ihre Ergebnisse mit den Tabellen auf:
[https://wikipedia.org/wiki/Letter_frequencies](https://wikipedia.org/wiki/Letter_frequencies)

*Tipp: Beachten Sie nur Kleinbuchstaben* <br>
*Tipp: Laden Sie die Dateien mit open(filename, 'r', encoding='utf8')*

Beispieltext Englisch: https://en.wikipedia.org/wiki/Lettuce <br>
Beispieltext Deutsch: https://de.wikipedia.org/wiki/Computer

<!-- grade_id: cell-cb779eb4d975d8f6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f4ef98edff3a5ccf, points: 1, grade: true, solution: false -->
```python
"""count_chars Tests"""
from nose.tools import assert_equal

# Test Englisch
fname = "files/texts/english.txt"
english_res = count_chars(fname)

assert_equal(23, len(english_res), "Die Anzahl der Buchstaben stimmt nicht")
assert_equal({'l': 89, 'e': 220, 't': 167, 'u': 68, 'c': 77, 'a': 164, 's': 124, 'i': 131, 'v': 23, 'n': 114, 'p': 29, 'o': 112, 'f': 38, 'h': 65, 'd': 73, 'y': 24, 'm': 46, 'r': 99, 'g': 33, 'w': 27, 'b': 21, 'k': 7, 'q': 2},
            english_res, "Die gezählten Werte stimmen nicht")

# Test Deutsch
fname = "files/texts/german.txt"
german_res = count_chars(fname)

assert_equal(25, len(german_res), "Die Anzahl der Buchstaben stimmt nicht")
assert_equal({'e': 335, 'i': 141, 'n': 211, 'c': 65, 'o': 66, 'm': 55, 'p': 34, 'u': 80, 't': 120, 'r': 153, 'g': 58, 'l': 70, 's': 112, 'h': 69, 'd': 76, 'a': 109, 'k': 18, 'j': 4, 'b': 40, 'v': 28, 'f': 26, 'z': 29, 'w': 22, 'y': 5, 'x': 1},
             german_res, "Die gezählten Werte stimmen nicht")
```